Before you begin, execute this cell to import numpy and packages from the D-Wave Ocean suite, and all necessary functions the gate-model framework you are going to use, whether that is the Forest SDK or Qiskit. In the case of Forest SDK, it also starts the qvm and quilc servers.

```python
%run -i "assignment_helper.py"
```

**Exercise 1** (1 point). This is a test exercise to ensure that you understand how a Vocareum submission works and what the workflow is for completing the coding assignments. Set the value 5 in a variable called `a`.

```python
#
# YOUR CODE HERE
#

```

The solution is evaluated in a cell with `assert` statements. Any exception raised in this cell will give you zero points for the exercise. Sometimes the assertions are hidden because they would give away the solution. Nevertheless, the evaluation side of the notebooks sees the content of the cell and will assign the points correctly. If the assertions are visible, you can ran the cell yourself to verify that your result is correct before submission. If you are stuck, don't be shy: start a discussion on the forum.

<!-- grade_id: exercise1, points: 1, grade: true, solution: false -->
```python
assert a == 5
```